```python
# Initialize Otter
import otter
grader = otter.Notebook("hw03.ipynb")
```

# Classes, inheritance, and exceptions
In this homework, we'll explore some of the subtleties of classes, inheritance, and execeptions, in our journey toward understanding how a reusable class is constructed. 

Let's construct a class that implements the "likes" relationship in FaceBook. Here's an example from a previous homework:  

```python
data = { 
     'Amy': {'Mike'}, 
     'John': {'Mike', 'Amy'}, 
     'Emily': {'Amy'},
     'George': {'Frank', 'Emily', 'Mae'}, 
     'Mae': {'John'}
}
from pprint import pprint
pprint(data)
```

Output:
```
{'Amy': {'Mike'},
 'Emily': {'Amy'},
 'George': {'Emily', 'Frank', 'Mae'},
 'John': {'Mike', 'Amy'},
 'Mae': {'John'}}
```

<!-- BEGIN QUESTION -->

Question 1: in the cell below, construct a class Likes that has the following specifications:  
1. A method likes(self, x, y) that returns `True` if x likes y, and `False` otherwise. 
2. A method add(self, x, y) that adds the fact "x likes y". 
3. A method remove(self, x, y) that removes the fact that "x likes y", if any such fact is present. 

```python
class Likes: 
    data = dict()
    def __init__(self, data = None): 
        if data is not None: 
            self.data = data
        else: 
            self.data = dict()
            
    def likes(self, x, y): 
        if x in self.data:
            if y in self.data[x]:
                return True
            else:
                return False
        else:
            return False
        
    def add(self, x, y): 
        if x in self.data:
            if not isinstance(self.data[x], list):
                self.data[x] = [self.data[x]]
            self.data[x].append(y)
        else:
            self.data[x] = y
            
    def remove(self, x, y): 
        if x in self.data:
            if y in self.data[x]:
                self.data[x].remove(y)
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

Question 2:  Create a subclass of `Likes` called `BetterLikes` that satisfies the following additional properties. 
1. If you give the constructor a non-conforming dictionary in the constructor, it raises a `TypeError` built-in exception. 
2. If you give any method a parameter that is not a string, it raises a `TypeError` exception. 
In other words, catch all type errors in methods. The point of this is that the user can't feed garbage to the class in any way. Otherwise, you can verify that garbage in input causes unpredictable behavior. 

```python
class BetterLikes(Likes):
    def __init__(self, data=None):
        if data is None: 
            data = dict()
        if type(data) is dict:
            for x in data:
                if type(data[x]) is not set:
                    raise TypeError
            super(BetterLikes, self).__init__(data) 
        else:
            raise TypeError
    def likes(self, x, y): 
        if type(x) is str and type(y) is str:
            return super(BetterLikes, self).likes(x, y) 
        else: 
            raise TypeError
    def add(self, x, y): 
        if type(x) is str and type(y) is str:
            super(BetterLikes, self).add(x, y) 
        else: 
            raise TypeError
        
    def remove(self, x, y): 
        if type(x) is str and type(y) is str:
            super(BetterLikes, self).remove(x, y)
        else: 
            raise TypeError
        
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

You are not done until you upload the exported zipfile to GradeScope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Output:
```
<IPython.core.display.HTML object>
```

Tests for q1 (points: 3)
```
>>> data = { 
...      'Amy': {'Mike'}, 
...      'John': {'Mike', 'Amy'}, 
...      'Emily': {'Amy'},
...      'George': {'Frank', 'Emily', 'Mae'}, 
...      'Mae': {'John'}}
>>> li = Likes(data)
>>> assert li.likes('Amy', 'Mike')
>>> assert not li.likes('Amy', 'George')
```
```
>>> li = Likes()
>>> li.add('Frank', 'Michelle')
>>> assert li.likes('Frank', 'Michelle')
```
```
>>> li = Likes()
>>> assert not li.likes('Frank', 'Michelle')
>>> assert not li.likes('Frank', 'George')
>>> li.add('Frank', 'Michelle')
>>> li.add('Frank', 'George')
>>> assert li.likes('Frank', 'George')
>>> assert li.likes('Frank', 'Michelle')
```
```
>>> data = { 
...      'Amy': {'Mike'}, 
...      'John': {'Mike', 'Amy'}, 
...      'Emily': {'Amy'},
...      'George': {'Frank', 'Emily', 'Mae'}, 
...      'Mae': {'John'}}
>>> li = Likes(data)
>>> assert li.likes('George', 'Mae')
>>> li.remove('George', 'Mae')
>>> assert not li.likes('George', 'Mae')
```
```
>>> li = Likes() 
>>> li.remove("George", "Mae")
```


Tests for q2 (points: 3)
```
>>> data = { 
...      'Amy': {'Mike'}, 
...      'John': {'Mike', 'Amy'}, 
...      'Emily': {'Amy'},
...      'George': {'Frank', 'Emily', 'Mae'}, 
...      'Mae': {'John'}}
>>> li = BetterLikes(data)
>>> assert li.likes('Amy', 'Mike')
>>> assert not li.likes('Amy', 'George')
```
```
>>> li = BetterLikes()
>>> li.add('Frank', 'Michelle')
>>> assert li.likes('Frank', 'Michelle')
```
```
>>> li = BetterLikes()
>>> assert not li.likes('Frank', 'Michelle')
>>> assert not li.likes('Frank', 'George')
>>> li.add('Frank', 'Michelle')
>>> li.add('Frank', 'George')
>>> assert li.likes('Frank', 'George')
>>> assert li.likes('Frank', 'Michelle')
```
```
>>> data = { 
...      'Amy': {'Mike'}, 
...      'John': {'Mike', 'Amy'}, 
...      'Emily': {'Amy'},
...      'George': {'Frank', 'Emily', 'Mae'}, 
...      'Mae': {'John'}}
>>> li = BetterLikes(data)
>>> assert li.likes('George', 'Mae')
>>> li.remove('George', 'Mae')
>>> assert not li.likes('George', 'Mae')
```
```
>>> try: 
...     foo = BetterLikes(set())
...     print("type error not caught")
... except TypeError: 
...     pass
```
```
>>> try: 
...     foo = BetterLikes({'amy': {'george': 1}})
...     print("type error not caught")
... except TypeError: 
...     pass
```
```
>>> foo = BetterLikes({})
>>> try: 
...     foo.add('George', set())
...     print("type error not caught")
... except TypeError: 
...     pass
```
```
>>> foo = BetterLikes({})
>>> try: 
...     foo.remove(set(), 'George')
...     print("type error not caught")
... except TypeError: 
...     pass
```
